```python
%pip install otter-grader
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

Output:
```
Collecting otter-grader
  Using cached otter_grader-6.1.6-py3-none-any.whl.metadata (3.9 kB)
Collecting click<9.0.0,>=8.1.7 (from otter-grader)
  Downloading click-8.3.1-py3-none-any.whl.metadata (2.6 kB)
Collecting dill>=0.3.0 (from otter-grader)
  Downloading dill-0.4.0-py3-none-any.whl.metadata (10 kB)
Collecting fica>=0.4.1 (from otter-grader)
  Using cached fica-0.4.1-py3-none-any.whl.metadata (2.0 kB)
Collecting ipylab<2.0.0,>=1.0.0 (from otter-grader)
  Using cached ipylab-1.1.0-py3-none-any.whl.metadata (6.7 kB)
Requirement already satisfied: ipython in /Users/arslan/Library/Python/3.12/lib/python/site-packages (from otter-grader) (8.23.0)
Collecting ipywidgets<9.0.0,>=8.1.5 (from otter-grader)
  Using cached ipywidgets-8.1.8-py3-none-any.whl.metadata (2.4 kB)
Collecting jinja2<4.0,>=3.1 (from otter-grader)
  Downloading jinja2-3.1.6-py3-none-any.whl.metadata (2.9 kB)
Collecting jupytext<2.0.0,>=1.16.4 (from otter-grader)
  Using cached jupytext-1.18.1-py3-none-any.whl.metadata
```

# Lab 1 – Python, NumPy, and Pandas


## Instructions

Welcome to the first assignment in DSC 259R!

This Jupyter Notebook contains the statements of the problems and provides code and Markdown cells to display your answers to the problems. The notebook is _only_ for displaying a readable version of your final answers. The coding will be done in an accompanying `lab.py` file that is imported into the current notebook, and **you will only submit that `lab.py` file**, not this notebook!

Some additional guidelines:

- **Labs will have both public tests and hidden tests.** The bulk of your grade will come from your scores on hidden tests, which you will only see on Gradescope after the assignment deadline.
- **Do not change the function names in the `lab.py` file!** The functions in the `lab.py` file are how your assignment is graded, and they are graded by their name. If you changed something you weren't supposed to, you can find the original code in the course GitHub repository.
- Notebooks are nice for testing and experimenting with different implementations before designing your function in your `lab.py` file. You can write code here, but make sure that all of your real work is in the `lab.py` file, since that's all you're submitting.
- You are encouraged to write your own additional helper functions to solve the lab, as long as they also end up in `lab.py`.

**To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.** More details on its usage are given at the bottom of this notebook.


**Importing code from `lab.py`**:

- Below, we import the `.py` file that's contained in the same directory as this notebook.
- We use the `autoreload` notebook extension to make changes to our `lab.py` file immediately available in our notebook. Without this extension, we would need to restart the notebook kernel to see any changes to `lab.py` in the notebook.
  - `autoreload` is necessary because, upon import, `lab.py` is compiled to bytecode (in the directory `__pycache__`). Subsequent imports of `lab` merely import the existing compiled python.


```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
from pathlib import Path
import io
import pandas as pd
import numpy as np
```

Let's get started! 🎉


## Part 1: Python Basics 🐍


### Question 0 – Consecutive Integers

Complete the implementation of the function `consecutive_ints`, which takes in a possibly empty list of integers (`ints`) and returns `True` if there exist two adjacent list elements that are consecutive integers and `False` otherwise.

For example, since `9` is next to `8`, `consecutive_ints([5, 3, 6, 4, 9, 8])` should evaluate to `True`, since `9` and `8` are consecutive integers. On the other hand, `consecutive_ints([1, 3, 5, 7, 9])` should evaluate to `False`.

**_Note_**: If you look at `lab.py`, you'll notice that the solution to this problem is already there. This question is done for you to show you what a completed homework problem looks like.


```python
# The cells below are here for you to write scratch work in.
# You should write the code for your answer in `lab.py`, not here.
```

```python
def consecutive_ints(ints):
    # We iterate up to len(ints) - 1 so we don't go out of bounds 
    # when checking the 'next' element
    for i in range(len(ints) - 1):
        # Check if the absolute difference between neighbors is 1
        if abs(ints[i] - ints[i+1]) == 1:
            return True
            
    # If the loop finishes without finding a pair, return False
    return False
    
```

```python

```

To run the public tests on your code for a given question, run the cell containing a call to `grader.check` that immediately follows it.

Remember, your grade will primarily be determined by hidden tests, which are **not** run when you run `grader.check`, so it's important to extensively test your functions on your own by calling them on different inputs. Does they work for edge cases? Real-world data is **very messy** and you should expect your data processing code to break without thorough testing!

You can write custom tests either by calling your functions on different inputs here in the notebook, or by writing doctests in `lab.py`.


```python
grader.check("q0")
```

Output:
```
q0 results: All test cases passed!
```

### Question 1 – Median vs. Mean

Complete the implementation of the function `median_vs_mean`, which takes in a non-empty list of numbers (`nums`) and returns `True` if median of the list is less than or equal to the mean of the list and `False` otherwise.

Recall, if a list has even length, the median is the mean of the middle two elements.

**_Note:_** In this question, you may only use built-in functions and methods in Python. You should not use `numpy` or `pandas` at all, nor should you import any additional packages.


```python
def median_vs_mean(nums):
    # Sort without modifying the original list
    s = sorted(nums)
    
    n = len(s)
    mid = n // 2
    
    # Compute median
    if n % 2 == 1:           # odd length
        median = s[mid]
    else:                    # even length
        median = (s[mid - 1] + s[mid]) / 2
    
    # Compute mean
    mean = sum(s) / n
    
    return median <= mean

```

```python

```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## Part 2: Strings and Files 🧵

The following questions will familiarize you with the basics of working with strings and reading data from files. Remember that by default, data from files are stored as strings in Python.


### Question 2 – $n$ Prefixes

Complete the implementation of the function `n_prefixes`, which takes a string `s` and a positive integer `n`. It returns a string containing the first `n` consecutive prefixes of `s` in reverse order.

For example, let's suppose `s` is the string `'Billy!'` and `n` is `4`. The consecutive prefixes of `'Billy!'` are:

- `'B'`
- `'Bi'`
- `'Bil'`
- `'Bill'`
- `'Billy'`
- `'Billy!'`

The first 4 of these are `'B'`, `'Bi'`, `'Bil'`, and `'Bill'`. If we combine these 4 in reverse order, we get `'BillBilBiB'`, which is what `n_prefixes('Billy!', 4)` should return. As another example, `n_prefixes('Marina', 3)` should return `'MarMaM'`. **You may assume that `n` is no larger than the length of `s`.**

**_Hint:_** Recall that [strings may be sliced](https://docs.python.org/3/tutorial/introduction.html#strings), like lists.


```python
def n_prefixes(s, n):
    result = ""
    
    # build prefixes from n down to 1
    for i in range(n, 0, -1):
        result += s[:i]
        
    return result

```

```python

```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

### Question 3 – Exploded Numbers 💣

Complete the implementation of the function `exploded_numbers`, which takes in a list of integers (`ints`) and a non-negative integer (`n`) and **returns a list of strings** containing numbers from the list expanded by `n` numbers in both directions, separated by spaces. Each integer should be [zero padded](https://www.tutorialspoint.com/python/string_zfill.htm) so that all integers outputted have the same length.

For example, consider `exploded_numbers([3, 8, 15], 2)`.

- If we explode 3 by 2 numbers in both directions, we get 1, 2, 3, 4, 5.
- If we explode 8 by 2 numbers in both directions, we get 6, 7, 8, 9, 10.
- If we explode 15 by 2 numbers in both directions, we get 13, 14, 15, 16, 17.

The longest length of any of the exploded numbers above is 2, so all of the outputted integers should have length 2.

- The string corresponding to 3 in the input is `'01 02 03 04 05'`.
- The string corresponding to 8 in the input is `'06 07 08 09 10'`.
- The string corresponding to 15 in the input is `'13 14 15 16 17'`.

So, `exploded_numbers([3, 8, 15], 2)` should return `['01 02 03 04 05', '06 07 08 09 10', '13 14 15 16 17']`.

As another example, `exploded_numbers([9, 99], 3)` should return `['006 007 008 009 010 011 012', '096 097 098 099 100 101 102']`.

**_Note_**: You can assume that negative numbers will never be encountered. That is, when testing your code, we will never explode a number so much that it becomes negative.


```python
def exploded_numbers(ints, n):
    exploded_lists = []
    
    # Step 1: build exploded ranges
    for x in ints:
        exploded = list(range(x - n, x + n + 1))
        exploded_lists.append(exploded)
    
    # Step 2: compute max width across all numbers
    max_width = 0
    for lst in exploded_lists:
        for num in lst:
            max_width = max(max_width, len(str(num)))
    
    # Step 3: format each exploded list as padded strings
    result = []
    for lst in exploded_lists:
        padded = [str(num).zfill(max_width) for num in lst]
        result.append(" ".join(padded))
    
    return result

```

```python

```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

### Question 4 – Reading Files

[Recall](https://docs.python.org/3/tutorial/inputoutput.html#reading-and-writing-files) that the built-in function `open` takes in a file path and returns _a file object_ (sometimes called a _file handle_). Below are a few properties of file objects:

- `open(path)` opens the file at location `path` for reading.
- `open(path)` is an _iterable_, which contains successive lines of the file.
- Once a file object is opened, after use it should be closed to avoid memory leaks. To ensure a file is closed once done, you should use a _context manager_ as follows:

```py
with open(path) as fh:
    for line in fh:
        process_line(line)
```

- To read the entire file into a string, use the `read` method:

```py
with open(path) as fh:
    s = fh.read()
```

However, you should be careful when reading an entire file into memory that the file isn't too big! _You should avoid this whenever possible!_

Complete the implementation of the function `last_chars`, which takes in file object (`fh`) and returns a string consisting of the last character of each line. Note that you don't have to use `open` at all; the argument given to you is a file object, not a file path.

**_Note:_** A newline (`'\n'`) is the "delimiter" of the lines of a file, and doesn't count as part of the line (as the tests imply). Every other character is part of the line. For more info on this, see [the interpretation](https://en.wikipedia.org/wiki/Newline#Interpretation) of files as a 'newline delimited variables' file.


```python
def last_chars(fh):
    result = ""
    
    for line in fh:
        line = line.rstrip('\n')   # remove only newline
        if line:                   # avoid errors if the line is empty
            result += line[-1]
            
    return result

```

```python

```

If your implementation is correct, you should see `'hrg'` when running the cell below:


```python
# You'll see the Path(...) / syntax a lot.
# It creates the correct path to your file,
# whether you're using Windows, macOS, or Linux.
# (Note that macOS and Linux use / to denote separate folders in paths,
# while Windows uses \.)

fp = Path("data") / "chars.txt"
last_chars(open(fp))
```

Output:
```
'hrg'
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## Part 3: `numpy` exercises 🥧

For a refresher on `numpy` and arrays, refer to the relevant section of the [these notes (from another UCSD course)](https://notes.dsc10.com/02-data_sets/arrays.html).


### Question 5 – Array Methods

Complete the implementations of the functions `add_root` and `where_square`. Specifications are given below. Your solutions should **not** contain any loops or list comprehensions.

#### `add_root`

`add_root` should take in a `numpy` array, `A`, and return a new `numpy` array that contains the element-wise sum of the elements in `A` with the _square roots of the positions of the elements in `A`_.

For instance, if `A` contains the values 5, 9, and 4, the output array should contain the values 5 (5 + $\sqrt{0}$), 10 (9 + $\sqrt{1}$), and 5.4142... (4 + $\sqrt{2}$).

<br>

#### `where_square`

`where_square` should take in a `numpy` array, `A`, and return a new `numpy` array of Booleans whose `i`th element is `True` if and only if the `i`th element of `A` is a perfect square.

For instance, `where_square(np.array([2, 9, 16, 15]))` should return `array([False, True, True, False])`.


```python
def add_root(A):
    return A + np.sqrt(np.arange(len(A)))
```

```python
def where_square(A):
    roots = np.sqrt(A)
    return roots == np.floor(roots)

```

```python
# Don't change this cell -- it is needed for the tests to work
A_1 = np.array([2, 4, 6, 7])
out_1 = add_root(A_1)

A_2 = np.array([1, 2, 16, 17, 32, 49])
out_2 = where_square(A_2)
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

### Question 6 - Filtering Matrices

Complete the implementation for the function `filter_cutoff_loop` and `filter_cutoff_np`.

#### Part 1: `filter_cutoff_loop`

`filter_cutoff_loop` should take in a matrix (2-dimensional `numpy` array) and a `float` cutoff. The function should return a 2-dimensional `numpy` array with only coumns that have a column mean (strictly) greater than the cutoff value. **This function should be implemented with loops or list comprehensions, not using `numpy` or `pandas` operations.** For example, given the matrix:

$$
a = \begin{bmatrix}
    1 & -2 & 3 & -3 \\
    0 & 1 & 3 & 6 \\
\end{bmatrix}
$$

The column means are computed as follows:

$$
\begin{bmatrix}
1\\
0
\end{bmatrix}
= \frac{1+0}{2} = 0.5 \text{, }
\begin{bmatrix}
-2\\
1
\end{bmatrix}
= \frac{-2+1}{2} = 0.5 \text{, }
\begin{bmatrix}
3\\
3
\end{bmatrix}
= \frac{3+3}{2} = 3 \text{, }
\begin{bmatrix}
-3\\
6
\end{bmatrix}
= \frac{-3+6}{2} = 1.5
$$

If the cutoff value is 1, only the last two columns have means above 1, thus only the last two columns are kept, as shown in the first example below.

```py
>>> a = np.array([[1,-2, 3,-3], [0, 1, 3, 6]])
>>> filter_cutoff_loop(a, 1.0)
# the column means are [0.1,-0.5, 3, 1.5]
# only columns 3 & 4 have column means higher than 1
# we only return columns 3 & 4
array([[ 3, -3],
       [ 3,  6]])
```

```py
>>> a = np.array([[1,-2, 3,-3], [0, 1, 3, 6]])
>>> filter_cutoff_loop(a, 2.0)
# the column means are [0.1,-0.5, 3, 1.5]
# only column 3 has a column mean higher than 2
# we only return column 3
array([[3],
       [3]])
```


```python
def filter_cutoff_loop(matrix, cutoff):
    rows = len(matrix)
    cols = len(matrix[0])

    keep_cols = []

    # Step 1: compute column means with loops
    for j in range(cols):
        col_sum = 0
        for i in range(rows):
            col_sum += matrix[i][j]
        mean = col_sum / rows

        if mean > cutoff:
            keep_cols.append(j)

    # Step 2: build the filtered matrix (still using loops)
    filtered = []
    for i in range(rows):
        row_vals = []
        for j in keep_cols:
            row_vals.append(matrix[i][j])
        filtered.append(row_vals)

    return np.array(filtered)
```

```python

```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

### Question 6

#### Part 2: `filter_cutoff_np`

`filter_cutoff_np` should take in the same inputs and output the same results as `filter_cutoff_loop`. **Do not use loops or list comprehensions for this implementation, only `numpy` operations.**

**_Hints:_**:

- What does the axis argument do in `np.mean`?
- Remember we can slice arrays using Boolean values. How does the code below work?

```py
>>> a = np.array([[1, 2, 3, 4], [5, 6, 7, 8]])
>>> a[:, [True, False, True, False]]
```


```python
import numpy as np

def filter_cutoff_np(matrix, cutoff):
    col_means = np.mean(matrix, axis=0)        # 1 × num_columns
    mask = col_means > cutoff                  # Boolean mask
    return matrix[:, mask]                     # select only columns we keep

```

```python

```

```python
grader.check("q6.2")
```

Output:
```
q6.2 results: All test cases passed!
```

### Question 7 – Stock Prices 📈

Complete the implementations of the functions `growth_rates` and `with_leftover`. Specifications are given below. Your solutions should **not** contain any loops or list comprehensions.

#### `growth_rates`

`growth_rates` should take in a `numpy` array, `A`, of [stock prices](https://en.wikipedia.org/wiki/Stock) for a single stock on successive days in USD. It should return an array of growth rates. That is, the `i`th number of the returned array should contain the rate of growth in stock price between the $i^{th}$ day to the $(i+1)^{th}$ day. The growth rate between two values is defined as $\frac{\text{final} - \text{initial}}{\text{initial}}$. You should return growth rates as **proportions, rounded to two decimal places**.

<br>

#### `with_leftover`

Again, suppose `A` is a `numpy` array of stock prices. Consider the following scheme:

- Suppose that you start each day with \$20 to purchase stocks.
- Each day, you purchase as many shares as possible of the stock. (The price changes each day, according to `A`.)
- Any money left-over after a given day is saved for possibly buying stock on a future day.

The function `with_leftover` should take in `A` and return the day (as an `int`) on which you can buy at least one full share using just "left-over" money. If this never happens, return `-1`. Note that the first stock purchase occurs on Day 0, and that you cannot purchase fractions of a share of a stock.

For example, if the stock price is \$3 every day, then the answer is `1` (corresponding to Day 1):

- Day 0: Buy 6 stocks with \\$20, and \\$2 is added to the leftover. Your total leftover is currently \\$2. This is not enough to buy one extra share, so you continue.
- Day 1: Buy 6 stocks with \\$20, and another \\$2 is added to the leftover. Your total leftover is now \\$4, so you can now buy one extra share. Hence, the answer is Day 1, and `with_leftover` should return `1`.

**_Hint:_** `np.cumsum` may be helpful.


```python
def growth_rates(A):
    # Differences between consecutive days
    diffs = A[1:] - A[:-1]
    # Elementwise growth rates
    rates = diffs / A[:-1]
    # Round to 2 decimal places
    return np.round(rates, 2)
```

```python
def with_leftover(A):
    # Money left over each day after buying as many whole shares as possible with $20
    # shares_per_day = floor(20 / price)
    # leftover_per_day = 20 - shares_per_day * price
    shares = np.floor(20 / A)
    leftovers = 20 - shares * A

    # Cumulative leftover over days
    cum_leftover = np.cumsum(leftovers)

    # Day where cumulative leftover is enough to buy 1 share at that day's price
    can_buy = cum_leftover >= A

    # Find first such day, if any
    indices = np.where(can_buy)[0]
    if indices.size == 0:
        return -1
    else:
        return int(indices[0])
```

```python
# Don't change this cell -- it is needed for the tests to work
fp = Path("data") / "stocks.csv"
stocks = np.array([float(x) for x in open(fp)])
out_3_stocks = growth_rates(stocks)

A_4 = np.array([3, 3, 3, 3])
out_4 = with_leftover(A_4)
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

## Part 4: Introduction to `pandas` 🐼

This part will help build familiarity with DataFrames in `pandas`.

As always for `pandas` questions:

1. Avoid writing loops through the rows of the DataFrame to do the problem, and
2. Test the output/correctness of your code with the help of the dataset given, but be sure your code will also run on data that is similar to but different from the dataset given. (One way to do this is to sample rows from the provided DataFrame using the `.sample` method).


The file `data/salary.csv` contains salary information for the 2021-22 National Basketball Association (NBA) season 🏀. Specifically, it contains the name, team, and salary of all players who have played at least 15 games last season. We will load this file and store it as a DataFrame named `salary`.


```python
# Do not edit this cell -- it is needed for the tests
salary_fp = Path("data") / "salary.csv"
salary = pd.read_csv(salary_fp)
salary.head()
```

Output:
```
Player Position           Team    Salary
0       John Collins       PF  Atlanta Hawks  23000000
1   Danilo Gallinari       PF  Atlanta Hawks  20475000
2  Bogdan Bogdanović       SG  Atlanta Hawks  18000000
3       Clint Capela        C  Atlanta Hawks  17103448
4       Delon Wright       SG  Atlanta Hawks   8526316
```

### Question 8 – `pandas` Basics

Your job is to complete the implementation of the function `salary_stats`, which takes in a DataFrame like `salary` and returns a **Series** containing the following statistics:

- `'num_players'`: The number of players.
- `'num_teams'`: The number of teams.
- `'total_salary'`: The total salary amount for all players.
- `'highest_salary'`: The name of the player with the highest salary. **Assume there are no ties.**
- `'avg_los'`: The average salary of the `'Los Angeles Lakers'`, rounded to two decimal places.
- `'fifth_lowest'`: The name and team of the player who has the fifth lowest salary, separated by a comma and a space (e.g. `'Billy Triton, Cleveland Cavaliers'`). **Assume there are no ties.**
- `'duplicates'`: A Boolean that is `True` if there are any duplicate last names, and `False` otherwise. Note that some players may have a suffix on their name, such as "Jr." or "III" -- you should ignore these. That is, "Billy Triton Jr." and "Tyler Triton" should be considered to have the same last name.
- `'total_highest'`: The total salary of the team that has the highest paid player.

The index of each element in the outputted Series is specified above.

**_Notes_**:

- Your function should work on a dataset of the same format that contains information from other years. This means that `salary_stats` should not "hard-code" any numbers or strings, but should compute them all programatically. In all cases, you may assume that none of the answers involving ranking involves a tie.
- The public tests don't test to see if your function actually returns the right numbers. You should manually inspect your result to make sure that all values seem appropriate.


```python
def salary_stats(salary):
    import pandas as pd  # safe even if already imported

    # Try to infer column names in a flexible way
    player_col = [c for c in salary.columns if 'player' in c.lower() or 'name' in c.lower()][0]
    team_col   = [c for c in salary.columns if 'team'   in c.lower()][0]
    salary_col = [c for c in salary.columns if 'sal'    in c.lower()][0]

    # 1. Number of players
    num_players = len(salary)

    # 2. Number of teams
    num_teams = salary[team_col].nunique()

    # 3. Total salary
    total_salary = salary[salary_col].sum()

    # 4. Name of player with highest salary (assume no ties)
    idx_max = salary[salary_col].idxmax()
    highest_salary_player = salary.loc[idx_max, player_col]

    # 5. Average salary of the Los Angeles Lakers, rounded to 2 decimals
    lakers_mask = salary[team_col] == 'Los Angeles Lakers'
    avg_los = round(salary.loc[lakers_mask, salary_col].mean(), 2)

    # 6. Player with 5th lowest salary: "Name, Team"
    fifth_row = salary.sort_values(by=salary_col, ascending=True).iloc[4]
    fifth_lowest = f"{fifth_row[player_col]}, {fifth_row[team_col]}"

    # 7. Duplicated last names (ignoring suffixes like Jr., III, etc.)
    suffixes = {'Jr.', 'Jr', 'Sr.', 'Sr', 'II', 'III', 'IV', 'V'}

    def last_name_only(full_name):
        parts = full_name.split()
        if len(parts) == 0:
            return ''
        if parts[-1] in suffixes and len(parts) >= 2:
            return parts[-2]
        return parts[-1]

    last_names = salary[player_col].apply(last_name_only)
    duplicates = last_names.duplicated().any()

    # 8. Total salary of the team that has the highest paid player
    team_of_highest = salary.loc[idx_max, team_col]
    total_highest = salary.loc[salary[team_col] == team_of_highest, salary_col].sum()

    return pd.Series(
        {
            'num_players':  num_players,
            'num_teams':    num_teams,
            'total_salary': total_salary,
            'highest_salary': highest_salary_player,
            'avg_los':      avg_los,
            'fifth_lowest': fifth_lowest,
            'duplicates':   duplicates,
            'total_highest': total_highest,
        }
    )
```

```python


```

```python
# Do not edit this cell -- it is needed for the tests
salary_fp = Path("data") / "salary.csv"
salary = pd.read_csv(salary_fp)
stats = salary_stats(salary)

salary_sample_fp = Path("data") / "salary_sample.csv"
salary_sample = pd.read_csv(salary_sample_fp)
sample_stats = salary_stats(salary_sample)
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

### Question 9 – Reading Malformed `.csv` Files

`data/malformed.csv` is a file of comma-separated values, containing the following fields:

| column name | description                                            | type    |
| ----------- | ------------------------------------------------------ | ------- |
| `'first'`   | first name of person                                   | `str`   |
| `'last'`    | last name of person                                    | `str`   |
| `'weight'`  | weight of person (lbs)                                 | `float` |
| `'height'`  | height of person (in)                                  | `float` |
| `'geo'`     | location of person; comma-separated latitude/longitude | `str`   |

Unfortunately, the entries contains errors with the placement of commas (`,`) and quotes (`"`) that cause `pandas`' `read_csv` function to fail parsing the file with the default settings. Don't believe us? Try using `pd.read_csv` on `data/malformed.csv` and look at what happens.

As a result, instead of using `pd.read_csv`, you must read in the file manually using Python's built-in `open` function.

Complete the implementation of the function `parse_malformed`, which takes in a file path (`fp`) and returns a parsed, properly-typed DataFrame with the information in the corresponding file. For example, `fp` may be `'data/malformed.csv'`. The DataFrame should contain the columns described in the data description table above (with the specified types).

**_Note:_**

- The only kinds of issues you need your function to handle are comma and quote misplacements; don't try and find any other issues with the CSV.
- With that said, you should assume that `data/malformed.csv` is a sample of a larger file that has the same sorts of errors, but potentially in different lines. For example, `data/malformed.csv` has an unnecessary quote `"` in line 4, but your function may be called on another CSV that has a perfectly fine line 4 but an unnecessary quote on some other line.
- So, **don't** implement `parse_malformed` assuming that the commas and quotes are mispositioned on specific lines; rather, implement `parse_malformed` such that it can handle these issues on every single line they appear in.
- A good way to proceed is to open `data/malformed.csv` and look carefully at the comma and quote placements.

The first few rows of `parse_malformed('data/malformed.csv')` should be:

<img src="./imgs/example-df.png" width=45%>


```python
import pandas as pd

def parse_malformed(fp):
    rows = []

    with open(fp) as fh:
        # Skip the header line
        header = fh.readline()

        for line in fh:
            line = line.strip()
            if not line:
                continue  # skip empty lines

            # Remove all double quotes
            line = line.replace('"', '')

            # Split by commas
            parts = line.split(',')

            if len(parts) < 4:
                # Malformed beyond simple comma/quote issues; skip or handle as needed
                continue

            first = parts[0].strip()
            last = parts[1].strip()

            # Find weight and height as the first two non-empty numeric tokens after last name
            weight = None
            height = None
            geo_parts = []

            for token in parts[2:]:
                tok = token.strip()
                if tok == '':
                    # skip empty tokens caused by extra commas
                    continue

                if weight is None:
                    weight = float(tok)
                elif height is None:
                    height = float(tok)
                else:
                    geo_parts.append(tok)

            # Join whatever remains as geo (may contain commas originally)
            geo = ",".join(geo_parts)

            rows.append({
                "first": first,
                "last": last,
                "weight": float(weight),
                "height": float(height),
                "geo": geo,
            })

    return pd.DataFrame(rows, columns=["first", "last", "weight", "height", "geo"])

```

```python
# Do not edit -- needed for tests
fp = Path("data") / "malformed.csv"
cols = ["first", "last", "weight", "height", "geo"]
df = parse_malformed(fp)
dg = pd.read_csv(fp, nrows=4, skiprows=10, names=cols)
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

## Congratulations! You're done with Lab 1! 🏁

As a reminder, all of the work you want to submit needs to be in `lab.py`.

To ensure that all of the work you want to submit is in `lab.py`, we've included a script named `lab-validation.py` in the lab folder. You shouldn't edit it, but instead, you should call it from the command line (e.g. the Terminal) to test your work.

Once you've finished the lab, you should open the command line and run, in the directory for this lab:

```
python lab-validation.py
```

**This will run all of the `grader.check` cells that you see in this notebook, but only using the code in `lab.py` – that is, it doesn't look at any of the code in this notebook. If all of your `grader.check` cells pass in this notebook but not all of them pass in your command line with the above command, then you likely have code in your notebook that isn't in your `lab.py`!**

You can also use `lab-validation.py` to test individual questions. For instance,

```
python lab-validation.py q1 q4 q7 q8
```

will run the `grader.check` cells for Questions 1, 4, 7, and 8 – again, only using the code in `lab.py`. The [video](#Infrastructure-Summary) linked above shows you how to use the script as well.

Once `python lab-validation.py` shows that you're passing all test cases, you're ready to submit your `lab.py` (and only your `lab.py`) to Gradescope. Once submitting to Gradescope, make sure to stick around until all test cases pass.

There is also a call to `grader.check_all()` below in _this_ notebook, but make sure to also follow the steps above.


```python

```

Tests for q0 (points: 0)
```
>>> consecutive_ints([5, 3, 6, 4, 9, 8])
True
```
```
>>> consecutive_ints([1, 3, 5, 7, 9]) == False
True
```
```
>>> consecutive_ints([]) == False
True
```


Tests for q1 (points: None)
```
>>> median_vs_mean([6, 5, 4, 3, 2])
True
```
```
>>> median_vs_mean([50, 20, 15, 40])
True
```
```
>>> median_vs_mean([1, 8, 9]) == False
True
```


Tests for q2 (points: None)
```
>>> n_prefixes('Billy', 4) == 'BillBilBiB'
True
```
```
>>> n_prefixes('Marina', 3) == 'MarMaM'
True
```
```
>>> n_prefixes('aaron', 2) == 'aaa'
True
```
```
>>> n_prefixes('Justin', 5) == 'JustiJustJusJuJ'
True
```


Tests for q3 (points: None)
```
>>> exploded_numbers([3, 8, 15], 2) == ['01 02 03 04 05', '06 07 08 09 10', '13 14 15 16 17']
True
```
```
>>> exploded_numbers([9, 99], 3) == ['006 007 008 009 010 011 012', '096 097 098 099 100 101 102']
True
```


Tests for q4 (points: None)
```
>>> fp = Path('data') / 'chars.txt'
>>> last_chars(open(fp)) == 'hrg'
True
```


Tests for q5 (points: None)
```
>>> isinstance(out_1, np.ndarray)
True
```
```
>>> bool(np.all(out_1 >= A_1))
True
```
```
>>> bool(np.isclose(out_1[3], 7 + np.sqrt(3)))
True
```
```
>>> isinstance(out_2, np.ndarray)
True
```
```
>>> out_2.dtype == np.dtype('bool')
True
```
```
>>> bool(out_2[2])
True
```


Tests for q6 (points: None)
```
>>> a = np.array([[1, -2, 3, -3], [0, 1, 3, 6]])
>>> len(filter_cutoff_loop(a, 1).shape) == 2
True
```
failure_message: should return a 2-dimensional numpy array
```
>>> a = np.array([[1, -2, 3, -3], [0, 1, 3, 6]])
>>> isinstance(filter_cutoff_loop(a, 1), np.ndarray)
True
```
failure_message: the data type returned should be a numpy array


Tests for q6.2 (points: None)
```
>>> a = np.array([[1, -2, 3, -3], [0, 1, 3, 6]])
>>> len(filter_cutoff_np(a, 1).shape) == 2
True
```
failure_message: should return a 2-dimensional numpy array
```
>>> a = np.array([[1, -2, 3, -3], [0, 1, 3, 6]])
>>> isinstance(filter_cutoff_np(a, 1), np.ndarray)
True
```
failure_message: the data type returned should be a numpy array


Tests for q7 (points: None)
```
>>> out_3_stocks.dtype == np.dtype('float')
True
```
```
>>> bool(out_3_stocks.max() == 0.03)
True
```
```
>>> import numbers
>>> isinstance(out_4, numbers.Integral)
True
```
```
>>> bool(out_4 == 1)
True
```


Tests for q8 (points: None)
```
>>> isinstance(stats, pd.Series)
True
```
```
>>> 'total_highest' in stats.index
True
```
```
>>> isinstance(stats.loc['duplicates'], bool) or isinstance(stats.loc['duplicates'], np.bool_)
True
```


Tests for q9 (points: None)
```
>>> list(df.columns) == cols
True
```
```
>>> df['last'].dtype == np.dtype('O')
True
```
```
>>> df['height'].dtype == np.dtype('float64')
True
```
```
>>> bool(df['geo'].str.contains(',').all())
True
```
```
>>> len(df) == 100
True
```
```
>>> dg.index = range(9, 13)
>>> bool((dg == df.iloc[9:13]).all().all())
True
```
failure_message: doctest examples
